---
title       : Data Visualization Principles, Part 2
description : Show the Data, Ease Comparisons - Use Common Axes, Consider Transformations, Ease Comparisons - Compared Visual Cues Should Be Adjacent
--- type:NormalExercise lang:r xp:100 skills:1 key:5ce63f8727
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
