---
title       : Data Visualization Principles, Part 3
description : Slope Charts, Encoding a Third Variable, Case Study - Vaccines, Avoid Pseudo and Gratuitous 3D Plots, Avoid Too Many Significant Digits

--- type:NormalExercise lang:r xp:100 skills:1 key:41c91866ca
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```