---
title       : Using the Gapminder Dataset
description : Faceting, Time Series Plots, Transformations, Stratify and Boxplot, Comparing Distributions, Density Plots, Ecological Fallacy
--- type:NormalExercise lang:r xp:100 skills:1 key:444ed32758
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
