---
title       : Summarizing with dplyr
description : dplyr, The Dot Placeholder, Group By, Sorting Data Tables

--- type:NormalExercise lang:r xp:100 skills:1 key:00ffbfee10
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
