---
title       : Introduction to ggplot2
description : Insert the chapter description here
--- type:NormalExercise lang:r xp:100 skills:1 key:de8c396d64
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
