---
title       : Customizing plots
description : Layers, Tinkering, Scales, Labels, and Colors, Add-on Packages, Other Examples

--- type:NormalExercise lang:r xp:100 skills:1 key:113a2cf0a8
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
