---
title       : Data Visualization Principles, Part 1
description : Introduction to Data Visualization Principles, Encoding Data Using Visual Cues, Know When to Include Zero, Do Not Distort Quantities, Order by a Meaningful Value

--- type:NormalExercise lang:r xp:100 skills:1 key:ba2e354cea
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
